```python
%matplotlib inline
```

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from nose.tools import *

from scipy.stats import ttest_ind

from sklearn.preprocessing import MinMaxScaler
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LogisticRegression
# Write your imports here
```

# Data Science Project Architecture Lab
## End-to-end project: demonstrating the power of OSEMN. Transition towards modelling. Machine learning basics

### Introduction
On 15 April 1912, the Titanic sank after colliding with an iceberg, killing more than two thirds of the crew and passengers. The dataset provided for you in the `data/` folder contains information about the passengers onboard and which of them survived.

The goal of this lab is to explore the data, prepare it for modelling, and perform a (kind of) simple classification. We'll also explore some basics of practical machine learning such as data preparation, testing and training sets, and model evaluation.

The original dataset is located [here](https://www.kaggle.com/c/titanic/data). You can read the page for more information about the data and variables.

### Problem 1. Read the dataset (1 point)
Read the dataset in the `data/titanic.csv` file into the `titanic_data` variable. Here's a short description of what each column means:
* PassengerId - a unique number identifying each passenger
* Survived - indicator variable: 1 if the passenger survived, 0 otherwise
* Pclass - passenger ticket class (1, 2 or 3). This can be used as an indicator of wealth
* Name
* Sex
* Age
* SibSp - number of siblings / spouses aboard the Titanic
* Parch - number of parents / children aboard the Titanic
* Ticket - ticket number
* Fare - passenger fare (price)
* Cabin - cabin number
* Embarked - port of embarkation: C = Cherbourg, Q = Queenstown, S = Southampton

**Notes on family relationships:**
* Sibling = brother, sister, stepbrother, stepsister
* Spouse = husband, wife (mistresses and fiancés were ignored)
* Parent = mother, father
* Child = daughter, son, stepdaughter, stepson. Some children travelled only with a nanny, therefore Parch = 0 for them.

Set the index column to be "PassengerId". Rename "Pclass" to "Class" and "Parch" to "ParCh". Other than that, the column names aren't too bad (although not Pythonic enough). Don't rename them.

Also, change the "Embarked" column to include the full names of the ports (see the column descriptions).

<!-- grade_id: cell-8d05536388210811, points: 0, grade: false, solution: true -->
```python
titanic_data = pd.read_csv("data/titanic.csv", index_col="PassengerId")
```

```python
titanic_data.rename(columns={'Pclass': 'Class', 'Parch': 'ParCh'}, inplace=True)
```

```python
def map_to_full_port_name(symbol):
    if symbol == "C":
        return "Cherbourg"
    elif symbol == "Q":
        return "Queenstown"
    elif symbol == "S":
        return "Southampton"
    else:
        return symbol
```

```python
titanic_data.Embarked = titanic_data.Embarked.apply(map_to_full_port_name)
```

```python
titanic_data.sample(n=30)
```

Output:
```
Survived  Class                                           Name  \
PassengerId                                                                   
573                 1      1               Flynn, Mr. John Irwin ("Irving")   
368                 1      3                 Moussa, Mrs. (Mantoura Boulos)   
448                 1      1                    Seward, Mr. Frederic Kimber   
689                 0      3                Fischer, Mr. Eberhard Thelander   
651                 0      3                              Mitkoff, Mr. Mito   
398                 0      2                        McKane, Mr. Peter David   
310                 1      1                 Francatelli, Miss. Laura Mabel   
152                 1      1              Pears, Mrs. Thomas (Edith Wearne)   
596                 0      3                    Van Impe, Mr. Jean Baptiste   
464                 0      2                   Milling, Mr. Jacob Christian   
729                 0      2                Bryhl, Mr. Kurt Arnol
```

<!-- grade_id: cell-eeefe71b639dffe8, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(titanic_data)
```

### Problem 2. Inspect and fill missing data (1 point)
See how many records are missing for each column. You can just execute the cell below.

```python
titanic_data.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
Int64Index: 891 entries, 1 to 891
Data columns (total 11 columns):
 #   Column    Non-Null Count  Dtype  
---  ------    --------------  -----  
 0   Survived  891 non-null    int64  
 1   Class     891 non-null    int64  
 2   Name      891 non-null    object 
 3   Sex       891 non-null    object 
 4   Age       714 non-null    float64
 5   SibSp     891 non-null    int64  
 6   ParCh     891 non-null    int64  
 7   Ticket    891 non-null    object 
 8   Fare      891 non-null    float64
 9   Cabin     204 non-null    object 
 10  Embarked  889 non-null    object 
dtypes: float64(2), int64(4), object(5)
memory usage: 83.5+ KB
```

You can see that most of the data is there. We have some people with unknown ages and two people with unknown embarkation ports.

For missing ages, there are three approaches. We can't say right now which will prove the most correct but we'll stick to one.
* Remove people with unknown ages - not desirable, since they are many
* Replace unknown ages with a "centinel" value, e.g. $-1$ - not desirable because this will introduce invalid data which may throw our models off 
* Replace unknown ages with the column mean

We'll stick with the third approach. Replace the `NaN` values in the `Age` column with the column mean.

<!-- grade_id: cell-cda0e1f62b8fa17e, points: 0, grade: false, solution: true -->
```python
round(titanic_data.Age.mean())
```

Output:
```
30
```

```python
mean_age = round(titanic_data.Age.mean())
titanic_data.Age = titanic_data.Age.fillna(mean_age)
```

```python
titanic_data.sample(n=20)
```

Output:
```
Survived  Class                                          Name  \
PassengerId                                                                  
322                 0      3                              Danoff, Mr. Yoto   
336                 0      3                            Denkoff, Mr. Mitto   
577                 1      2                          Garside, Miss. Ethel   
568                 0      3   Palsson, Mrs. Nils (Alma Cornelia Berglund)   
352                 0      1        Williams-Lambert, Mr. Fletcher Fellows   
739                 0      3                            Ivanoff, Mr. Kanio   
884                 0      2                 Banfield, Mr. Frederick James   
244                 0      3                 Maenpaa, Mr. Matti Alexanteri   
652                 1      2                           Doling, Miss. Elsie   
881                 1      2  Shelley, Mrs. William (Imanita Parrish Hall)   
315                 0      2                            Hart, Mr. Benjamin   
```

Let's inspect missing embarkation ports. Store the passengers with unknown embarkation ports in the provided variable.

<!-- grade_id: cell-c81adf03dbc34dba, points: 0, grade: false, solution: true -->
```python
passengers_with_unknown_embarkation_ports = titanic_data[titanic_data.Embarked.isnull()]
passengers_with_unknown_embarkation_ports
```

Output:
```
Survived  Class                                       Name  \
PassengerId                                                               
62                  1      1                        Icard, Miss. Amelie   
830                 1      1  Stone, Mrs. George Nelson (Martha Evelyn)   

                Sex   Age  SibSp  ParCh  Ticket  Fare Cabin Embarked  
PassengerId                                                           
62           female  38.0      0      0  113572  80.0   B28      NaN  
830          female  62.0      0      0  113572  80.0   B28      NaN
```

We can see there are two such passengers with the same ticket. We can check there are no other passengers with the same ticket number. We have no idea what to do but we might just replace them with the most common embarkation port.

Find out which port was the most common. Replace the two NaN values in the dataset with this port.

<!-- grade_id: cell-bd2f821dd9cb5fc9, points: 0, grade: false, solution: true -->
```python
most_common_port = titanic_data.groupby("Embarked").Embarked.count().idxmax()
```

```python
titanic_data.Embarked = titanic_data.Embarked.fillna(most_common_port)
```

<!-- grade_id: cell-50f02a8a39bf9d82, points: 1, grade: true, solution: false -->
```python
# Tests for all operations
assert_false(titanic_data.Age.isnull().any())

assert_is_not_none(passengers_with_unknown_embarkation_ports)
assert_is_not_none(most_common_port)
assert_false(titanic_data.Embarked.isnull().any())
```

### Problem 3. Remove unnecessary columns (1 point)
The `Cabin` column contains too many missing values. Probably the best we can do with it is remove it. Also, the names and ticket numbers might be useful in another analysis, but not in this case. We're interested in which passengers survived and we have no reason to think that their names might be related to their survival rate. Also, the ticket numbers are somewhat random.

**Note:** It might be interesting to extract the titles of the passengers (e.g. "Mr.", "Miss", "Dr.", etc.) and see whether it correlates to survival rate (e.g. people with higher social status might be more likely to get a boat and survive). But let's not focus on this right now. The class and ticket fare are good enough to indicate social status / wealth.

Remove the `Cabin`, `Name`, and `Ticket` columns.

<!-- grade_id: cell-4d96142a29f5f032, points: 0, grade: false, solution: true -->
```python
titanic_data = titanic_data.drop(columns=["Cabin", "Name", "Ticket"])
```

```python
titanic_data.head()
```

Output:
```
Survived  Class     Sex   Age  SibSp  ParCh     Fare     Embarked
PassengerId                                                                   
1                   0      3    male  22.0      1      0   7.2500  Southampton
2                   1      1  female  38.0      1      0  71.2833    Cherbourg
3                   1      3  female  26.0      0      0   7.9250  Southampton
4                   1      1  female  35.0      1      0  53.1000  Southampton
5                   0      3    male  35.0      0      0   8.0500  Southampton
```

<!-- grade_id: cell-ce236d7fee7f5854, points: 1, grade: true, solution: false -->
```python
assert_equal(titanic_data.shape, (891, 8))
```

### Problem 4. Explore the data: single variables (1 point)
Let's start with visualizing single variables. 

Try plotting a histogram of all ages with 20 bins. You'll see a kind of unusual peak. Remember that this is because we filled in the missing data with the mean of all ages, and it happens to be right where that peak is.

Also, try plotting a bar chart (or a pie chart) showing the number of passengers who are male and female. To do this, group the dataset by sex and count the number of rows for each group. `num_passengers_by_sex` should be a `pd.Series` with  two indices: "male" and "female".

Finally, try plotting a histogram of fares to see how asymmetric they are.

**Note:** The plots are not autograded, only the data. Feel free to change them, experiment, and add more plots as you see fit. I had quite a lot of fun playing around with different aspects of the data. This is the reason to have EDA, after all :).

**Note 2:** The variables should be really simple to set.

<!-- grade_id: cell-2c3caaa38c49514a, points: 0, grade: false, solution: true -->
```python
all_ages = None
num_passengers_by_sex = None
all_fares = None

# YOUR CODE HERE
raise NotImplementedError()

plt.hist(all_ages, bins = 20)
plt.title("Distribution of ages")
plt.show()

plt.gca().set_aspect("equal")
plt.pie(num_passengers_by_sex, labels = num_passengers_by_sex.index, autopct = "%.2f%%")
plt.title("Passengers per sex")
plt.show()

plt.hist(all_fares, bins = 20)
plt.title("Distribution of fares")
plt.show()
```

<!-- grade_id: cell-1e84086c6a0454d7, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(all_ages)
assert_is_not_none(num_passengers_by_sex)
assert_is_not_none(all_fares)

assert_equal(len(all_ages), len(all_fares))
assert_equal(num_passengers_by_sex.index.tolist(), ["female", "male"])
```

### Problem 5. Explore correlations in the dataset (1 point)
We can play a lot with single variables, groups, etc. But let's focus on correlations now.

One of the first things we can do is check all correlations on all variables, like so:

```python
titanic_data.corr()
```

Well, there are some correlations but it seems nothing too interesting can be found.

Let's now try some groupings. For example, what percentage of each gender survived? Recall that we calculated the total number of passengers for each gender in the previous exercise.

Filter the `titanic_data` dataset to get only survived passengers and apply the same grouping and counting as you did in the previous exercise. You should get a series with "male" and "female" as the indices.

If your answers are correct, the `print()` statements should run without errors.

<!-- grade_id: cell-7554388a9c07ce6f, points: 0, grade: false, solution: true -->
```python
survived_passengers = None
# YOUR CODE HERE
raise NotImplementedError()
print("Survived men: {} / {}, {:.2f}%".format(survived_passengers.male, num_passengers_by_sex.male, survived_passengers.male / num_passengers_by_sex.male * 100))
print("Survived women: {} / {}, {:.2f}%".format(survived_passengers.female, num_passengers_by_sex.female, survived_passengers.female / num_passengers_by_sex.female * 100))
```

We can see that a far greater proportion of women survived. This is really significant for two reasons: 1) the difference is really large (74% women vs. 19% men survived), 2) the total number of women on board is smaller.

We can therefore conclude that women have been given advantage while evacuating from the ship.

<!-- grade_id: cell-508e9ba1aadd8279, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(survived_passengers)
assert_equal(num_passengers_by_sex.index.tolist(), ["female", "male"])
```

Feel free to look for more correlations if you wish.

Let's now focus on something else: the distribution of ages broken down by class. As we already mentioned, passenger class can be used as a proxy for a person's wealth.

Group the dataset by class and extract the ages for each group. Store this in the `ages_by_class` variable. It should be a `pd.Series` with `Class` as the index.

Plot a histogram showing the three age distributions.

<!-- grade_id: cell-b0f18ef015029cc9, points: 0, grade: false, solution: true -->
```python
ages_by_class = None
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Note: this is not an autograded cell. It's here only to help you 
# find out whether your answer and data format are correct
assert_is_not_none(ages_by_class)
assert_equal(ages_by_class.size().tolist(), [216, 184, 491])
```

```python
for passenger_class, ages in ages_by_class:
    plt.hist(ages, label = "Class {}".format(passenger_class), alpha = 0.7)
plt.title("Distribution of passenger ages per class")
plt.legend()
plt.show()
```

We can see something really interesting. It seems that first-class passengers are a little bit older than third-class passengers. But is this really true? We can't tell for sure. First of all, there are many more third-class passengers; and second, we can't be sure whether there's a significant difference or not.

Fortunately, there's a rigorous statistical method to find out. Enter **hypothesis testing**.

### Problem 6. Perform hypothesis testing on age vs. class (1 point)
First, let's store "class 1" and "class 3" passenger ages in their own variables, for easier work.

```python
first_class_ages = ages_by_class.get_group(1)
third_class_ages = ages_by_class.get_group(3)
```

To perform a hypothesis test, we'll need a hypothesis. Actually, a pair of hypotheses. The "null hypothesis", $H_0$ says that "there's nothing interesting going on with the data". The "alternative hypothesis", $H_1$ says the opposite.

We want to prove whether or not the passenger class is correlated with the age. Therefore:
* $H_0:$ Passenger class is not correlated with age. `first_class_ages` and `third_class_ages` are two samples from the same distribution.
* $H_1:$ `first_class_ages` and `third_class_ages` come from two different distributions.

Ideally, **we'd like to reject the null hypothesis**.

Here's a quick explanation of the process: we'll perform a test. The exact details aren't important. We assume that $H_0$ is true, therefore **the differences between the two histograms occur simply by chance**. The test will return a $p$-value. It corresponds to the probability that we observe **as extreme or more extreme differences** between the two histograms if $H_0$ is really true.

We have to agree on a "threshold value" of $p$. Usually that's 5% (0.05), but let's choose 1% in this case. What does this mean? If we reject $H_0$, there will still be 1% chance that we rejected it wrongly.

**If $p\le1\%$, we will reject $H_0$**.

To compare the two variables, it's easiest to perform what's called a **t-test**. It's already been imported for you. Call it like this: `test_result = ttest_ind(<first_variable>, <second_variable>, equal_var = False)`.

**Note:** You can get additional information about the mechanics of statistical hypothesis testing on the Internet. Research more if you wish. You can also research what `equal_var = False` is and why we aren't allowed to assume equal variances in this case.

<!-- grade_id: cell-2fad87583bb70604, points: 0, grade: false, solution: true -->
```python
test_result = None
# YOUR CODE HERE
raise NotImplementedError()
print(test_result.pvalue)
if test_result.pvalue <= 0.01:
    print("The differences in age are significant. Reject H0.")
else:
    print("There's not enough evidence to reject H0. Don't accept or reject anything else.")
```

<!-- grade_id: cell-7ce2e934e8d8ecb9, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(test_result)
```

Therefore, we can conclude that **the distributions of ages are significantly different at the 1% level**. Actually, they're so significantly different, that we might as well have chosen $1.10^{-17}\%$ and still be correct.

This means that ages are different for the different classes. How can we interpret this? Probably wealthier people are older. Younger people might not need, or might not be able to afford, a higher fare.

### Problem 7. Prepare the data for modelling: indicator variables (1 point)
We're going to use `scikit-learn` to model the data. However, that's not so simple. We first need to preprocess the data a little.

Most importantly, all variables should be numeric. `scikit-learn` doesn't know how to deal with text and categories.

We need to convert `Sex` and `Embarked` to categories. There are many ways to do that.

What's considered the best way is via the so-called "indicator variables". These are variables whose values are 0 or 1. For example, let's look at the "Sex" column. It has two possible values: "male" and "female". Each of these values will create a new column: `Sex_male` and `Sex_female`. If the passenger is male, he will have 1 in the `Sex_male` column, and so on. Similarly, with `Embarked`.

There's a really easy way to do this in `pandas`: `pd.get_dummies(dataframe)`. Note that this returns another dataframe. Add the columns: `["Class", "Sex", "Embarked"]` to the dataframe. Write the code and explore the newly created dataset.

<!-- grade_id: cell-ce48a98dc2da3cce, points: 0, grade: false, solution: true -->
```python
titanic_data_for_modelling = None
# YOUR CODE HERE
raise NotImplementedError()
titanic_data_for_modelling.head()
```

We can now see that we have more columns. We can also see that since `Sex` has only two possible values, the two columns `Sex_female` and `Sex_male` are just opposites of each other. We can safely remove one of them. However, this is not true for the `Class` and `Embarked` columns.

<!-- grade_id: cell-eb71a6ae99067bb1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cf11c03fee894a9c, points: 1, grade: true, solution: false -->
```python
assert_equal(titanic_data_for_modelling.shape, (891, 12))
```

Later, it will be really convenient to separate the explanatory variables from the target variable.

We want to predict whether or not a person has survived. Therefore, `Survived` will be our target variable. All other variables will be our explanatory variables (also called features).

```python
titanic_data_features = titanic_data_for_modelling.drop("Survived", axis = 1)
titanic_data_target = titanic_data_for_modelling.Survived
```

### Problem 8. Prepare the data for modelling: normalization (1 point)
In order for the model to perform better, we usually need to rescale the values for each numeric column. 

Why do we do this? It's related to algorithm stability and convergence. Generally, a machine learning algorithm will perform better if all values are in similar ranges.

Do we always need it? No, but we usually do.

There are many types of normalization. In this case, we're going to use a **min-max normalization**. The minimum value in the column will become 0, the maximum will become 1. All values in between will be scaled accordingly.

`scikit-learn` has a very convenient [MinMaxScaler](http://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.MinMaxScaler.html#sklearn.preprocessing.MinMaxScaler). You use it by simply instantiating it and passing the data:
```python
scaler = MinMaxScaler()
titanic_data_features_scaled = scaler.fit_transform(titanic_data_features)
```

Note that `titanic_data_scaled` will be a 2D `numpy` array.

<!-- grade_id: cell-15df658fa2716e19, points: 0, grade: false, solution: true -->
```python
titanic_data_scaled = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8290e1e8c30f5922, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(titanic_data_features_scaled)
assert_equal(titanic_data_features_scaled.shape, (891, 11))
```

### Problem 9. Split the data (0 points)
When we want to evaluate a machine learning model, we usually hide some data from it. We train the model on most of the data, but when we test it afterwards, we pass the additional, hidden data. This is similar to how humans learn - a teacher won't give the exact answers to all students. If this was the case, the teacher cannot know whether a student really learned something, or just memorized all the answers.

The function `train_test_split` from `scikit-learn` will perform the splitting for us. See the docs [here](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html).

We usually want $\approx 70\%$ of the data for training and the remaining $\approx 30\%$ for testing. It's very important that the data is shuffled. `train_test_split()` will do this by default.

We'll pass the features and target variables and we'll get the different parts accordingly.

```python
features_train, features_test, target_train, target_test = train_test_split(
    titanic_data_features_scaled, titanic_data_target, train_size = 0.7, test_size = 0.3, random_state = 42)
print(features_train.shape, features_test.shape, target_train.shape, target_test.shape, sep = "\r\n")
```

### Problem 10. Model the data (1 point)
Let's model the data using logistic regression. That's very simple.

First, create a logistic regression model (with no custom settings). Then, fit the model using the training features and training target.
```python
model = LogisticRegression()
model.fit(???, ???)
```

If you wish, you can inspect the model coefficients and intercept.

<!-- grade_id: cell-f0dd6abc403dec0a, points: 0, grade: false, solution: true -->
```python
model = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4afb23030f0e02b5, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(model)
```

### Problem 11. Score the model (1 point)
The default scoring metric for a classification model is **accuracy**. Use `model.score(???, ???)` to get an accuracy score for the model. This should be around 80%.

**Note:** Remember to use `features_test` and `target_test`, not the training subsets.

<!-- grade_id: cell-528f747f698aeadf, points: 0, grade: false, solution: true -->
```python
score = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0e64bb27e29b0292, points: 1, grade: true, solution: false -->
```python
assert_greater(score, 0)
assert_less_equal(score, 1)
```

You see that even though it might seem difficult at first, working with models is pretty easy.

Feature preparation, train / test split, normalization, extraction of explanatory features vs. target, modelling, testing, and evaluating: these are all parts of the data modelling process. It's the basic idea of **machine learning**.

We started from a dataset and we were able to explore, visualize, and model the data. After all this, we have several deliverables: notebook with our research, model (that we might upload somewhere - but that's outside the scope of this lab), (mostly) repeatable research. We have followed a careful and complete process to get to the final results.

We can, of course, extend the study. But this is enough for now.

# Good luck on the exam! :)